# How to complete an assignment:
    
- in each exercise exchange the ```# YOUR CODE HERE raise NotImplemented Error()``` with your solution
- as you create the solution run the tests frequently to see if they pass
- remember you can write more tests if you'd like (I recommend it).
- your solution MUST BE in the cell you are replacing a placeholder. If you are writing other functions to help yourself they also MUST BE in that cell. Otherwise automatic marking script will not give you points. (you'll notice when you use 'Validate' described below).
- when you are done, in your Noteable you can "Validate" the assignment (it will basically run all the tests)
- then when you are ready to submit, in your Noteable choose "Submit". If you by accident submit many times, it's fine. The last submition before the deadline will count.

### Each of the 10 tasks below is worth 10 points. Total is 100 points.

(i.e. each set of tests will give you 5 points if they pass). 

### Solved example

<!-- grade_id: cell-df0e92ef19aab615, points: 0, grade: false, solution: true -->
```python
# to see it working, replace the body of the function with code: return number * 2 

def double_the_number(number):
    """
    function will return double of the number given as an argument
    """
    ### BEGIN SOLUTION
    return 2 * number
    ### END SOLUTION
```

<!-- grade_id: cell-6f00d799be290e2e, points: 0, grade: true, solution: false -->
```python
assert double_the_number(4) == 8
assert double_the_number(-2) == -4
assert double_the_number(0) == 0
print("all tests passed")
```

### Task 1

hint: you can use if elif < <= > >= and or

<!-- grade_id: cell-8cb073f43a2e5d7b, points: 0, grade: false, solution: true -->
```python
def find_smallest_number(n1,n2,n3,n4):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    if n1 <= n2 and n1 <= n3 and n1 <= n4:
        return n1
    elif n2 <= n1 and n2 <= n3 and n2 <= n4:
        return n2
    elif n3 <= n1 and n3 <= n2 and n3 <= n4:
        return n3
    else:
        return n4

    ### END SOLUTION
```

```python
# remember: you can always just run the function to see its outcome, like this:
find_smallest_number(8,4,5,6)
```

<!-- grade_id: cell-450086c09907e437, points: 5, grade: true, solution: false -->
```python
assert find_smallest_number(2,1,3,4) == 1
assert find_smallest_number(-1,0,1,3) == -1
### BEGIN HIDDEN TESTS
assert find_smallest_number(2,6,1,4) == 1
assert find_smallest_number(2,1,6,-4) == -4
### END HIDDEN TESTS
```

<!-- grade_id: cell-762fa40f992456ae, points: 5, grade: true, solution: false -->
```python
assert find_smallest_number(1,2,1,4) == 1
assert find_smallest_number(-1,-1,-1,-1) == -1
### BEGIN HIDDEN TESTS
assert find_smallest_number(1,-2,-2,-2) == -2
assert find_smallest_number(0,0,0,0) == 0
### END HIDDEN TESTS
```

### Task 2

<!-- grade_id: cell-c89a06f8ec1bb5c1, points: 0, grade: false, solution: true -->
```python
def are_all_numbers_not_zero(n1,n2,n3,n4):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    return n1 != 0 and n2 != 0 and n3 != 0 and n4 != 0 
    ### END SOLUTION
```

```python
are_all_numbers_not_zero(1,2,3,4)
```

<!-- grade_id: cell-5e8bccf5527edda8, points: 5, grade: true, solution: false -->
```python
assert are_all_numbers_not_zero(1,2,3,0) == False
assert are_all_numbers_not_zero(1,2,3,4) == True
### BEGIN HIDDEN TESTS
assert are_all_numbers_not_zero(1,0,1,1) == False
assert are_all_numbers_not_zero(1,1,-1,1) == True

### END HIDDEN TESTS
```

<!-- grade_id: cell-8ed3430f91888b65, points: 5, grade: true, solution: false -->
```python
assert are_all_numbers_not_zero(1,2,0,0) == False
assert are_all_numbers_not_zero(0,0,0,0) == False
### BEGIN HIDDEN TESTS
assert are_all_numbers_not_zero(1,0,1,0) == False
assert are_all_numbers_not_zero(0,1,0,1) == False
### END HIDDEN TESTS
```

### Task 3

New element: to get a length of a word we use the len function: `len( the_word )` or `len( "banana" )` 

```python
# here are some examples:
print( len("banana") )
print( len("banana") * 3)

length_of_plum = len("plum")
print(length_of_plum)
```

<!-- grade_id: cell-5394a718088c0279, points: 0, grade: false, solution: true -->
```python
def is_word_shorter_than_10_characters(word):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    return len(word) < 10
    ### END SOLUTION
```

<!-- grade_id: cell-30208478d64d8ef2, points: 5, grade: true, solution: false -->
```python
assert is_word_shorter_than_10_characters("") == True
assert is_word_shorter_than_10_characters("banana") == True
assert is_word_shorter_than_10_characters("all the bananas") == False
### BEGIN HIDDEN TESTS
assert is_word_shorter_than_10_characters("banana") == True
assert is_word_shorter_than_10_characters("a massive fruit bowl") == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-6a61cd3ffbef2bb5, points: 5, grade: true, solution: false -->
```python
assert is_word_shorter_than_10_characters("ripe plum") == True
assert is_word_shorter_than_10_characters("ripe plums") == False
assert is_word_shorter_than_10_characters("a ripe plum") == False

### BEGIN HIDDEN TESTS
assert is_word_shorter_than_10_characters("grapefrui") == True
assert is_word_shorter_than_10_characters("grapefruit") == False
assert is_word_shorter_than_10_characters("grapefruits") == False
### END HIDDEN TESTS
```

### Task 4


<!-- grade_id: cell-3074008799e7785c, points: 0, grade: false, solution: true -->
```python
def the_longer_word(w1,w2):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    if len(w1) >= len(w2):
        return w1
    else:
        return w2
    ### END SOLUTION
```

```python
the_longer_word("pineapple", "apple")
```

<!-- grade_id: cell-de4ccacb1d9dbcc0, points: 5, grade: true, solution: false -->
```python
assert the_longer_word("cats", "cat") == "cats"
assert the_longer_word("dog", "dogs") == "dogs"
### BEGIN HIDDEN TESTS
assert the_longer_word("a", "") == "a"
assert the_longer_word("a", "bb") == "bb"
### END HIDDEN TESTS
```

<!-- grade_id: cell-35b509a1a0759c40, points: 5, grade: true, solution: false -->
```python
assert the_longer_word("cat", "cat") == "cat"
assert the_longer_word("", "") == ""
### BEGIN HIDDEN TESTS
assert the_longer_word("a", "a") == "a"
assert the_longer_word("123456789", "123456789") == "123456789"

### END HIDDEN TESTS
```

### Task 5


<!-- grade_id: cell-006c4e5b8c8e8da7, points: 0, grade: false, solution: true -->
```python
def the_longest_word(w1,w2,w3,w4):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    longest = w1
    if len(w2) >= len(longest):
        longest = w2
    if len(w3) >= len(longest):
        longest = w3
    if len(w4) >= len(longest):
        longest = w4
        
    return longest
    ### END SOLUTION
```

```python
the_longest_word("a","aa","aaaa","aaa")
```

<!-- grade_id: cell-ecc1e610f22d42a3, points: 5, grade: true, solution: false -->
```python
assert the_longest_word("a","aa","aaaa","aaa") == "aaaa"
assert the_longest_word("aaaaa","","aaa","aaa") == "aaaaa"
### BEGIN HIDDEN TESTS
assert the_longest_word("a","a","a","aaa") == "aaa"
assert the_longest_word("a","aaa","a","a") == "aaa"
### END HIDDEN TESTS
```

<!-- grade_id: cell-18d69f26554e8143, points: 5, grade: true, solution: false -->
```python
assert the_longest_word("aaaaa","aa","aaaaa","aaa") == "aaaaa"
assert the_longest_word("a","a","a","a") == "a"
### BEGIN HIDDEN TESTS
assert the_longest_word("aaa","aaa","aaa","aaa") == "aaa"
assert the_longest_word("","","","") == ""
### END HIDDEN TESTS
```

### Task 6

<!-- grade_id: cell-34a9f3680a585373, points: 0, grade: false, solution: true -->
```python
def total_length_of_words(w1,w2,w3,w4):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    return len(w1) + len(w2) + len(w3) + len(w4)
    ### END SOLUTION
```

```python
total_length_of_words("aa","aaa","aaaa","")
```

<!-- grade_id: cell-2fbaebbe56677b0b, points: 5, grade: true, solution: false -->
```python
assert total_length_of_words("a","aa","aaaa","aaa") == 10
assert total_length_of_words("aaaaaa","aa","aa","a") == 11
### BEGIN HIDDEN TESTS
assert total_length_of_words("a","a","a","aaa") == 6
assert total_length_of_words("a","a","a","a") == 4
assert total_length_of_words("bananabananabanana","bananabananabanana","bananabananabanana","bananabananabanana") == 72
### END HIDDEN TESTS
```

<!-- grade_id: cell-891f9fa9ddfdfe50, points: 5, grade: true, solution: false -->
```python
assert total_length_of_words("a","","aaaa","aaa") == 8
assert total_length_of_words("","","","a") == 1
### BEGIN HIDDEN TESTS
assert total_length_of_words("a","","","a") == 2
assert total_length_of_words("","","","") == 0
assert total_length_of_words("","aa","","aa") == 4
### END HIDDEN TESTS
```

### Task 7

<!-- grade_id: cell-ba76ea05b733027a, points: 0, grade: false, solution: true -->
```python
def is_word_longer_than(word, some_length):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    return len(word) > some_length
    ### END SOLUTION
```

<!-- grade_id: cell-72d4f5e339baadda, points: 5, grade: true, solution: false -->
```python
assert is_word_longer_than("banana", 5) == True
assert is_word_longer_than("banana", 7) == False
### BEGIN HIDDEN TESTS
assert is_word_longer_than("plums", 4) == True
assert is_word_longer_than("plums", 6) == False
### END HIDDEN TESTS
```

<!-- grade_id: cell-b5aebb5a20db19b0, points: 5, grade: true, solution: false -->
```python
assert is_word_longer_than("banana", 6) == False
### BEGIN HIDDEN TESTS
assert is_word_longer_than("plums", 5) == False
### END HIDDEN TESTS
```

### Task 8

<!-- grade_id: cell-dd61a0e5ef03253d, points: 0, grade: false, solution: true -->
```python
def is_word_lengths_between_lengths(word, mininum_length, maximum_length):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    return len(word) >= mininum_length and len(word) <= maximum_length 
    ### END SOLUTION
```

<!-- grade_id: cell-4793e277c3cd3bc6, points: 5, grade: true, solution: false -->
```python
assert is_word_lengths_between_lengths("banana", 3,7) == True
assert is_word_lengths_between_lengths("plum", 8,12) == False
### BEGIN HIDDEN TESTS
assert is_word_lengths_between_lengths("banana", 9,12) == False
assert is_word_lengths_between_lengths("plum", 2,8) == True
### END HIDDEN TESTS
```

<!-- grade_id: cell-4bd5d6e74f428965, points: 5, grade: true, solution: false -->
```python
assert is_word_lengths_between_lengths("banana", 3,6) == True
assert is_word_lengths_between_lengths("banana", 6,9) == True
assert is_word_lengths_between_lengths("banana", 7,9) == False
assert is_word_lengths_between_lengths("banana", 4,5) == False
### BEGIN HIDDEN TESTS
assert is_word_lengths_between_lengths("plum", 2,4) == True
assert is_word_lengths_between_lengths("plum", 4,8) == True
assert is_word_lengths_between_lengths("plum", 2,3) == False
assert is_word_lengths_between_lengths("plum", 5,8) == False
### END HIDDEN TESTS
```

### Task 9 & 10 - Challenge

You can solve this by writing a few functions, one calling another, but remember to write them all inside of the same cell with the solution.

#### Adding Strings:

In this exercise we will edit strings. We'll learn how to do it properly in the incomming weeks, but for now, just add them like

```
name = "Pim"
greeting = "Hello " + name + ", how are you?"
```

to achieve ```"Hello Pim, how are you?"```. Notice that spaces need to be included inside of the brackets if you'd like to keep them (eg. that space right after `Hello` but still inside of the bracket).

#### Adding Strings and Integers

Also: remember that when you add a string to a number, you need to makr sure that Python will know if you mean a number or a string. This is called Casring.

```
age = 12
greeting = "I am " + str(age) + " years old"
```


**Important Note: Keep below functions in one cell, this will enable you to use some functions in other functions without confusing the grade marking engine.**

<!-- grade_id: cell-756eeff6e85a660c, points: 0, grade: false, solution: true -->
```python
def report_about_the_day(week_day_as_word):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    day_number  = which_day_of_the_week(week_day_as_word)
    return "It is "+which_day_with_ending(day_number)+" day of the week. "+day_number_into_weekend_status(day_number)
    ### END SOLUTION
    
def which_day_of_the_week(week_day):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    if week_day == "Monday":
        return 1
    elif week_day == "Tuesday":
        return 2
    elif week_day == "Wednesday":
        return 3
    elif week_day == "Thursday":
        return 4
    elif week_day == "Friday":
        return 5
    elif week_day == "Saturday":
        return 6
    elif week_day == "Sunday":
        return 7
    ### END SOLUTION
    
def which_day_with_ending(day_number):
        ### BEGIN SOLUTION
    if day_number == 1:
        return "1st"
    elif day_number == 2:
        return "2nd"
    elif day_number == 3:
        return "3rd"
    else:
        return str(day_number)+ "th"
        ### END SOLUTION
    
def day_number_into_weekend_status(day_number):
    """
    here in your own words describe what this function will do and how
    """
    ### BEGIN SOLUTION
    days_till_weekend = 6 - day_number
    if days_till_weekend <= 0:
        return "Weekend is now"
    elif days_till_weekend == 1:
        return "Weekend is tomorrow"
    else:
        return str(days_till_weekend) + " days till weekend"

    ### END SOLUTION

    
```

<!-- grade_id: cell-261ab844ca1e5728, points: 5, grade: true, solution: false -->
```python
assert which_day_of_the_week("Monday") == 1
assert which_day_of_the_week("Friday") == 5

### BEGIN HIDDEN TESTS
assert which_day_of_the_week("Thursday") == 4
assert which_day_of_the_week("Sunday") == 7
### END HIDDEN TESTS
```

<!-- grade_id: cell-559f23b6b3418a9a, points: 5, grade: true, solution: false -->
```python
assert which_day_with_ending(1) == "1st"
assert which_day_with_ending(2) == "2nd"
assert which_day_with_ending(3) == "3rd"
assert which_day_with_ending(4) == "4th"

### BEGIN HIDDEN TESTS
assert which_day_with_ending(5) == "5th"
assert which_day_with_ending(6) == "6th"

### END HIDDEN TESTS
```

<!-- grade_id: cell-045e2488d28c9d2b, points: 5, grade: true, solution: false -->
```python
assert day_number_into_weekend_status(1) == "5 days till weekend"
assert day_number_into_weekend_status(5) == "Weekend is tomorrow"
assert day_number_into_weekend_status(6) == "Weekend is now"

### BEGIN HIDDEN TESTS
assert day_number_into_weekend_status(3) == "3 days till weekend"
assert day_number_into_weekend_status(7) == "Weekend is now"
### END HIDDEN TESTS
```

<!-- grade_id: cell-46264c68f67c56cc, points: 5, grade: true, solution: false -->
```python
assert report_about_the_day("Monday") == "It is 1st day of the week. 5 days till weekend"
assert report_about_the_day("Tuesday") == "It is 2nd day of the week. 4 days till weekend"
assert report_about_the_day("Wednesday") == "It is 3rd day of the week. 3 days till weekend"
assert report_about_the_day("Friday") == "It is 5th day of the week. Weekend is tomorrow"
assert report_about_the_day("Saturday") == "It is 6th day of the week. Weekend is now"
assert report_about_the_day("Sunday") == "It is 7th day of the week. Weekend is now"
### BEGIN HIDDEN TESTS
assert report_about_the_day("Thursday") == "It is 4th day of the week. 2 days till weekend"
### END HIDDEN TESTS
```

```python

```