Please do not complete this lab by yourself before your timetabled lab time. These notebooks are designed for pairs and if you complete it alone, you loose away peer-learning opportunity from you and your partner. Thanks!

# Parts of this lab:

1. Examples of Class Diagrams and how to write code from diagrams - read these with your partner. Take turns being the main reader.
2. Drawing Task - use whiteboard with your partner
3. Coding Tasks: you are given diagrams and tests - just write the solution code (take turns, pair-program)


# UML Class Diagrams

We will focus on planning simple business solutions using the code you already know. You will draw your solutions on paper first and then for some of them you will write some code.
**CLASS DIAGRAMS in UML (Universal Modeling Language)**

Below you will be creating a set of class diagrams. Like you've seen during the lecture, a class diagram consists of three boxes on top of each other

- top box is the class name
- second box is class variables: their names and type of values they hold
- third box is class functions: their name, arguments they take (name and type), and return value

| **Class_name** |
|:--------|
|variable_name: Variable_type   |
|-------------------------------------------|
|method ( argument_1, argument_2, ... ) ---> return_type  |

These are a tool for you to communicate with other programmers, so feel free to add anything else that is relevant for you.

# Three Examples of diagrams and code that was build based on them:

As you look at these examples think if you could have written the code, if you were given the diagrams (and maybe some tests)

| **Book** |
|:---|
|title: String |
|author: String |
|unique_book_number: Integer |
|-------------------------------------------|
|text_on_the_cover( ) ---> String  |

```python
class Book:
    def __init__(self, title, author, number):
        self.title = title
        self.author = author
        self.unique_book_number = number
        
    def text_on_the_cover( self ):
        return "{} by {}".format(self.title, self.author)
    
a_book = Book("About Bananas", "Dr Fruit", 123456)
print(a_book.text_on_the_cover()) 

assert a_book.text_on_the_cover() == "About Bananas by Dr Fruit"
```

| **Student** |
|:---|
|name: String |
|student_number: Integer |
|loaned_books: [Book] |
|-------------------------------------------|
|number_of_loaned_books( ) ---> int  |
|loan_a_new_book( Book ) |
|return_all_books( ) ---> [Book]  |

```python
class Student:
    def __init__(self, name, student_number):
        self.name = name
        self.student_number = student_number
        self.loaned_books = []
      
    def number_of_loaned_books(self):
        return len(self.loaned_books)
        
    def loan_a_new_book(self, new_book ):
        self.loaned_books.append(new_book)
    
    def return_all_books(self ):
        self.loaned_books.clear()

a_student = Student("Mia Banks", 103333 )
book1 = Book("About Bananas", "Dr Fruit", 123456)
book2 = Book("About Pinapples", "Dr Fruit", 123457)
print( a_student.number_of_loaned_books() )

a_student.loan_a_new_book(book1)
a_student.loan_a_new_book(book2)
print( a_student.number_of_loaned_books() )

a_student.return_all_books()
print( a_student.number_of_loaned_books() )
```

| **Library** |
|:---|
|books: [Book] |
|-------------------------------------------|
|add_new_book( Book )  |
|add_new_books( [Book] )  |
|number_of_books( ) ---> Integer  |
|get_book_with_number( Integer ) ---> Book  |
|get_books_of_author( String ) ---> [Book]  |
|loan_book_with_number_to_student( Integet, Student ) |
|student_returns_all_books( Student ) |

```python
class Library:
    def __init__(self):
        self.books = []
    
    def add_new_book(self, new_book ):
        self.books.append(new_book)
    
    def add_new_books(self, new_books ):
        self.books.extend(new_books)

    def number_of_books(self ):
        return len(self.books)
    
    def get_book_with_number(self, searched_book_number ):
        for book in self.books:
            if book.unique_book_number == searched_book_number:
                removed_book = self.books.remove(book)   
                return removed_book
    
    def get_books_of_author(self, searched_author ):
        searched_books = []
        for book in self.books:
            if book.author == searched_author:
                searched_books.append(book)
        return searched_books

      
```

```python
a_library = Library()
print(a_library.number_of_books())

book1 = Book("About Bananas", "Dr Fruit", 123456)
book2 = Book("About Pinapples", "Dr Fruit", 123457)
a_library.add_new_book(book1)
a_library.add_new_book(book2)
print(a_library.number_of_books())

book_number_123456 = a_library.get_book_with_number(123456)
print(a_library.number_of_books())

a_student = Student("Mia Banks", 103333 )
a_student.loan_a_new_book(book_number_123456)
print( a_student.number_of_loaned_books() )
```

# Drawing Task  (no code yet) 

**DIAGRAM TWO REAL WORLD SCENARIOS**. 

- Pick TWO scenarios of real-world objects interacting (see examples below for inspiration)
- Draw UML diagrams for each scenario. It should not be more complex than 2-3 classes per scenario
- In each class. do not exceed 3-4 variables and 3-4 methods

- **Look at the clock - do not spend more than 5 minutes on diagramming each scenario**

note: It's better to do simple class diagrams but many times, rather than complex ones but just once.

Some inspiration, but feel free to come up with your own examples. Just remember: no more than 5 mins per diagram.

- Food truck next to the Business School: possible classes Drink, Meals, Client, Truck, Ingredient...
- Netflix: possible classes Movie, Actor, User, Playlist
- JustEat food delivery: possible classes Restaurant, Client, Delivery, Meal, DeliveryDriver
- Your own examples...

```python
# do not write any code for this (yet), just do it on the whiteboard. Remember to take a screemshot when you are done
```

# Coding Tasks:

You are given diagrams and tests - just write the solution code (take turns, pair-program)

Note: When numbers have decimal places, the trailing zeros (zeros at the end) will be trimmed (not displayed). That means that when you print 1.20000 you will see only 1.2 - it does not mean that your result is inacurate, but only that meaningless zeros at the end are not displayed.

So ```1.2 == 1.2000000``` is True, and tests like ```assert 1.2 == 1.20``` will pass. Meaningless zeros will be removed, but that does not mean the number is rounded or somehow changed.

### TASK 1

Here's a Class Diagram for a fizzy drink Can, write a class based on it. You will be given some tests

| **Can** |
|:---|
|volume: Float |
|brand: String |
|flavour: String |
|-------------------------------------------|
|init( Float, String, String) # (volume, brand, flavour) |
|will_it_fit_in_glass_sized(Float) ---> Boolean # return true volume is smaller than given glass size|

<!-- grade_id: cell-c8c89af04ac473f8, points: 0, grade: false, solution: true -->
```python
# here write your Can class

class Can:
    # your code here
```

<!-- grade_id: cell-d198d883bc9c5c5b, points: 1, grade: true, solution: false -->
```python
coke = Can(0.33, "Coke", "Diet")
assert coke.volume == 0.33
assert coke.brand == "Coke"
assert coke.flavour == "Diet"
# hidden tests are here
### BEGIN HIDDEN TESTS
fanta = Can(0.5, "Fanta", "Grape")
assert fanta.volume == 0.5
assert fanta.brand == "Fanta"
assert fanta.flavour == "Grape"
### END HIDDEN TESTS
```

<!-- grade_id: cell-b0893a9dd46e7506, points: 1, grade: true, solution: false -->
```python
coke = Can(0.33, "Coke", "Diet")
assert coke.will_it_fit_in_glass_sized(0.10) == False
assert coke.will_it_fit_in_glass_sized(0.33) == True
assert coke.will_it_fit_in_glass_sized(0.60) == True
# hidden tests are here
### BEGIN HIDDEN TESTS
fanta = Can(0.5, "Fanta", "Grape")
assert fanta.will_it_fit_in_glass_sized(0.10) == False
assert fanta.will_it_fit_in_glass_sized(0.50) == True
assert fanta.will_it_fit_in_glass_sized(0.60) == True
### END HIDDEN TESTS
```

### TASK 2

Here's a Class Diagram for a cinema ticket, write a class based on it. You will be given some tests

| **Ticket** |
|:---|
|movie_title: String |
|seat: String |
|screen_number: Number |
|price: Float |
|time: String # for now let's just use String for time, like "16:10" |
|-------------------------------------------|
|init( String, String, Integer, Float, String, String)   |
|is_it_for_this_seat_and_screen( String, Integer ) ---> Boolean  # ( seat, screen ) eg. returns True/False whether seat and screen are correct  |
|as_string_simple( ) ---> Float  # eg.  "Screen 4 Seat F7" *** |
|as_string_all_info( ) ---> Float  # eg.  "Titanic - 16:10 - Screen 4 Seat F7 - £6.20"  *** |
|can_i_afford_it_for( Float ) ---> Boolean  # eg. returns True/False describing if price is lower or equal to the Float argument |

Note: parts marked with *** are harder and for full points

<!-- grade_id: cell-97f95d2452bd5623, points: 0, grade: false, solution: true -->
```python
# here write your Ticket class

class Ticket:
    # your code here
```

<!-- grade_id: cell-eccdd4d002ec01a7, points: 1, grade: true, solution: false -->
```python
ticket1 = Ticket(movie_title="Titanic", seat="B7", screen_number=3, price=6.20, time="17:00") 
assert ticket1.movie_title == "Titanic"
assert ticket1.seat == "B7"
assert ticket1.screen_number == 3
assert ticket1.price == 6.20
assert ticket1.time == "17:00"
# hidden tests are here
### BEGIN HIDDEN TESTS
ticket2 = Ticket(movie_title="Titanic 2", seat="B8", screen_number=4, price=7.20, time="18:00") 
assert ticket2.movie_title == "Titanic 2"
assert ticket2.seat == "B8"
assert ticket2.screen_number == 4
assert ticket2.price == 7.20
assert ticket2.time == "18:00"
### END HIDDEN TESTS
```

<!-- grade_id: cell-8f8020e349759c88, points: 1, grade: true, solution: false -->
```python
ticket1 = Ticket(movie_title="Titanic", seat="B7", screen_number=3, price=6.20, time="17:00")

assert ticket1.is_it_for_this_seat_and_screen("B7", 3) == True
assert ticket1.is_it_for_this_seat_and_screen("G2", 3) == False
assert ticket1.is_it_for_this_seat_and_screen("B7", 6) == False
assert ticket1.is_it_for_this_seat_and_screen("G2", 6) == False

# hidden tests are here
### BEGIN HIDDEN TESTS
ticket2 = Ticket(movie_title="Titanic 2", seat="B8", screen_number=4, price=7.20, time="18:00") 

assert ticket2.is_it_for_this_seat_and_screen("B8", 4) == True
assert ticket2.is_it_for_this_seat_and_screen("B8", 3) == False
assert ticket2.is_it_for_this_seat_and_screen("B7", 4) == False
assert ticket2.is_it_for_this_seat_and_screen("G2", 6) == False
### END HIDDEN TESTS

```

<!-- grade_id: cell-cd419b93a62aed7c, points: 1, grade: true, solution: false -->
```python
ticket1 = Ticket(movie_title="Titanic", seat="B7", screen_number=3, price=6.20, time="17:00")
assert ticket1.can_i_afford_it_for(5.10) == False
assert ticket1.can_i_afford_it_for(6.20) == True
assert ticket1.can_i_afford_it_for(10.00) == True

# hidden tests are here
### BEGIN HIDDEN TESTS
ticket2 = Ticket(movie_title="Titanic 2", seat="B8", screen_number=4, price=7.20, time="18:00") 

assert ticket2.can_i_afford_it_for(5.10) == False
assert ticket2.can_i_afford_it_for(7.20) == True
assert ticket2.can_i_afford_it_for(10.00) == True
### END HIDDEN TESTS

```

<!-- grade_id: cell-364545f1b241f1c3, points: 1, grade: true, solution: false -->
```python
ticket1 = Ticket(movie_title="Titanic", seat="B7", screen_number=3, price=6.20, time="17:00")
assert ticket1.as_string_simple() == "Screen 3 Seat B7"

# hidden tests are here
### BEGIN HIDDEN TESTS
ticket2 = Ticket(movie_title="Titanic", seat="B8", screen_number=4, price=6.20, time="17:00")
assert ticket2.as_string_simple() == "Screen 4 Seat B8"
### END HIDDEN TESTS


assert ticket1.as_string_all_info() == "Titanic - 17:00 - Screen 3 Seat B7 - £6.20"
```

<!-- grade_id: cell-e043127dae2389a7, points: 1, grade: true, solution: false -->
```python
ticket1 = Ticket(movie_title="Titanic", seat="B7", screen_number=3, price=6.20, time="17:00")
assert ticket1.as_string_all_info() == "Titanic - 17:00 - Screen 3 Seat B7 - £6.20"

# hidden tests are here
### BEGIN HIDDEN TESTS
ticket2 = Ticket(movie_title="Titanic 2", seat="C9", screen_number=4, price=6.20, time="19:00")
assert ticket2.as_string_all_info() == "Titanic 2 - 19:00 - Screen 4 Seat C9 - £6.20"
### END HIDDEN TESTS
```

### TASK 3

Here's a Class Diagram simmilar to the one you've seen during the lab, write a class based on it. You will be given some tests

| **Product** |
|:---|
|name: String |
|production_cost: Float |
|markup_percent: Float | # with value 0.15 for 15% markup
|bulk_size: Integer # when we buy this many OR MORE then bulk_percent_discount will be applied  # | 
|bulk_percent_discount: Float # price discounted by that much if we buy more items than bulk size # |
|-------------------------------------------|
|init( String, Float, Float, Integer, Float)  # name, cost, markup_percent, bulk_size, bulk_percent_discount # |
|profit_per_item( ) ---> Float  # production_cost * markup_percent # |
|shop_cost_of_item( ) ---> Float  # production_cost * (1.0 + markup_percent) # |
|how_many_can_client_buy_at_normal_price( Float ) ---> Int  # how many can we buy for this much # *** |
|does_buying_this_many_give_me_discount( Int ) ---> Boolean |
|how_many_can_client_buy_at_if_discount_applies( Float ) ---> Int  # only apply dicount if enough items are sold # *** 
|profit_to_be_made_if_client_buys_for( Float ) ---> Float  # what will be the profit if I buy and sell this many # *** 

Note: parts marked with *** are harder 

hint: add `__repr__` function to make debugging easier 

<!-- grade_id: cell-c89a06f8ec1bb5c1, points: 0, grade: false, solution: true -->
```python
# here wrote your Product class

    
class Product:
    # your code here
```

```python
# add more cells and create some tests of your own

beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 

# here's some prints to get you started
print(beans.shop_cost_of_item())
print(beans.profit_per_item())
print(beans.how_many_can_i_buy_for(5.0))
print(beans.profit_to_be_made_if_i_buy_for(5.0))
print(beans.how_many_can_i_buy_for(20.0))
print(beans.profit_to_be_made_if_i_buy_for(20.0))
```

<!-- grade_id: cell-ab98f2fd69dfd7a1, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.name == "Beans"
assert beans.production_cost == 1.00
assert beans.markup_percent == 0.20
assert beans.bulk_size == 8
assert beans.bulk_percent_discount == 0.10
        
# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.name == "Bananas"
assert product.production_cost == 2.00
assert product.markup_percent == 0.30
assert product.bulk_size == 12
assert product.bulk_percent_discount == 0.30
### END HIDDEN TESTS
```

<!-- grade_id: cell-e752112b115d6ce4, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.shop_cost_of_item() == 1.20

# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.shop_cost_of_item() == 2.6
### END HIDDEN TESTS
```

<!-- grade_id: cell-d6e79a354bc3c544, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.profit_per_item() == 0.20

# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.profit_per_item() == 0.6
### END HIDDEN TESTS
```

<!-- grade_id: cell-192833752cfb958f, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.how_many_can_i_buy_for(5.0) == 5 # notice this is a whole number
# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.how_many_can_i_buy_for(5.0) == 2 # notice this is a whole number
### END HIDDEN TESTS
```

<!-- grade_id: cell-28f341792d83852c, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.how_many_can_i_buy_for(20.0) == 22 # notice this is a whole number
# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.how_many_can_i_buy_for(20.0) == 14 # notice this is a whole number
### END HIDDEN TESTS
```

<!-- grade_id: cell-e752115bb62d6ce4, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.profit_to_be_made_if_i_buy_for(5.0) == 1.0
# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.profit_to_be_made_if_i_buy_for(5.0) == 1.2
### END HIDDEN TESTS
```

<!-- grade_id: cell-e752112b112d6ce4, points: 1, grade: true, solution: false -->
```python
beans = Product(name="Beans", production_cost=1.00, markup_percent=0.20, bulk_size=8, bulk_percent_discount=0.10) 
assert beans.profit_to_be_made_if_i_buy_for(10.0) == 2.2
# hidden tests are here
### BEGIN HIDDEN TESTS
product = Product(name="Bananas", production_cost=2.00, markup_percent=0.30, bulk_size=12, bulk_percent_discount=0.30) 
assert product.profit_to_be_made_if_i_buy_for(20.0) == 8.4
### END HIDDEN TESTS
```

# If you made it this far: here's a treat!

Pick a scenario from your drawing task on the beginning and write code for classes for that scenario. 

```python

```